```python
# Initialize Otter
import otter
grader = otter.Notebook("hw10_sp24.ipynb")
```

# Homework 10: Linear Regression

**Helpful Resource:**

- [Python Reference](http://data8.org/sp22/python-reference.html): Cheat sheet of helpful array & table methods used in Data 8!

**Recommended Readings**: 

* [The Regression Line](https://www.inferentialthinking.com/chapters/15/2/Regression_Line.html)
* [Method of Least Squares](https://www.inferentialthinking.com/chapters/15/3/Method_of_Least_Squares.html)
* [Least Squares Regression](https://www.inferentialthinking.com/chapters/15/4/Least_Squares_Regression.html)

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute this cell again.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. **Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Note: This homework has hidden tests on it. That means even though the tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**


Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

You should start early so that you have time to get help if you're stuck.

```python
# Run this cell to set up the notebook, but please don't change it.

import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
from datetime import datetime
```

## 1. Free Throws vs. 3 Point Shots

Does having a skill in one area of basketball shooting imply skill in another? While both [free throws](https://en.wikipedia.org/wiki/Free_throw) and [three-point shots](https://en.wikipedia.org/wiki/Three-point_field_goal) require precision and technique, they differ significantly in their execution. This exercise investigates whether players who excel at free throws also tend to perform well at three-point shots.
Since we're learning about linear regression, we will look specifically for a linear association between skill level in the two types of shots.


The dataset includes the 2023-2024 free throw percentages (FT %) and three-point percentages (3PT %) for 50 NBA players. We have selected players with the highest free throw percentages to examine if this skill correlates with a high three-point shooting percentage. In the NBA, a three-point percentage over 35% is considered good, and over 40% is considered excellent.





```python
# Run this cell to load the data
shots = Table.read_table('three_vs_freethrow.csv').select("FT %","3PT %")
shots
```

**Question 1.1.** Create a function `std_units` that converts the values in the array `data` to standard units. **(5 points)**


```python
def std_units(data):
    ...
```

```python
grader.check("q1_1")
```

**Question 1.2.** Using the `std_units` function created previously, define a new function named `correlation` which computes the correlation coefficient between two arrays, `x` and `y`. **(5 points)**

```python
def correlation(x, y):
    ...
```

```python
grader.check("q1_2")
```

<!-- BEGIN QUESTION -->

**Question 1.3.** Visual analysis of data through plotting is a crucial step in statistical analysis as it helps identify patterns, trends, and potential outliers. Effective visualization can provide immediate insights that are not always obvious through numerical methods alone.

Task: Create a scatter plot **that includes the regression line**, with the free throw percentages on the horizontal axis and the three-point percentages on vertical axis. **(5 points)**

See the documentation on `scatter` [here](http://data8.org/datascience/_autosummary/datascience.tables.Table.scatter.html#datascience.tables.Table.scatter) for instructions on how to have Python draw the regression line automatically.

*Hint:* The `fit_line` argument may be useful here!

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.4.** By examining the scatter plot provided, with free throw percentages (FT %)  on the horizontal axis and three point (3PT %) on the vertical axis, estimate the correlation coefficient r. Does r appear to be closest to 0, 0.5, or -0.5? Justify your answer with an explanation based on the trend observed in the plot.

 **(5 points)**


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 1.5.** In order to understand the relationship between two basketball statistics, such as free throw percentages (FT%) and three-point percentages (3PT %), it's useful to calculate key regression statistics that summarize this relationship.

Develop a function named `compute_regression_stats` that takes in the argument `tbl`, a data table with two columns. The first column is the x-axis (the explanatory variable), and the second column is the y-axis (the response variable). The function should return an array with three elements: **(1) The correlation coefficient** of the two columns, and the **(2) slope** and **(3) intercept** of the regression line that predicts the second column from the first, in original units. **(5 points)**

*Hint:* This is a rare occasion where it’s better to implement the function using column indices instead of column names, in order to be able to call this function on any table. If you need a reminder about how to use column indices to pull out individual columns, please refer to [this](https://www.inferentialthinking.com/chapters/06/Tables.html#accessing-the-data-in-a-column) section of the textbook.

```python
def compute_regression_stats(tbl):
    ...
    return make_array(r, slope, intercept)

constraints = compute_regression_stats(shots) 
print('r:', constraints.item(0), '; slope:', constraints.item(1), '; intercept:', constraints.item(2))
```

```python
grader.check("q1_5")
```

**Question 1.6.**: Now suppose you want to go the other way and predict a free throw percentage (FT %) given a three-point percentage (3PT %). Now 3PT % becomes the explanatory variable used to predict FT % the new response variable. How would various regression statistics be affected? Set `regression_changes` equal to an array that includes the numbers of the correct answer choices below. **(5 points)**

1) The correlation coefficient will remain unchanged because it only measures the strength and direction of the *linear* relationship between the two variables, regardless of which is the predictor.

2) The value of the slope will remain the same because it only reflects the rate of change between the two variables.

3) The value of the intercept may change because it is calculated based on the distribution and scale of the predictor variable.

4) The units of the slope will not change because the units of the predictor and response variables are independent of each other.

```python
regression_change = ...
regression_change
```

```python
grader.check("q1_6")
```

**Question 1.7.** Let's use `compute_regression_stats` (from Question 1.5) to predict what certain NBA Players' three point percentage (3PT %) would be given their Free Throw percentage (FT %). **(5 points)**

The NBA player with the highest free throw percentage (FT %) is Klay Thompson with 92.7%. What is the prediction for Thompson's three point perecentage (3PT %) using this line?



```python
free_throw_3pt_est = ...
print("Predicted free throw percentage: {:.1f}%".format(np.round(free_throw_3pt_est, 1)))
```

```python
grader.check("q1_7")
```

<!-- BEGIN QUESTION -->

**Question 1.8.** Is it reliable to use this line to predict three point percentage (3PT %) of players? **(5 points)**

*Hint:* Compare Thompson' three point percentage (3PT %)  to the three point percentage in `shots`. Is it relatively similar to the rest of the data (shown in Question 1.3)? 


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Smartphone Sales

Imagine you're a business analyst in December 2020. Smartphones are essential consumer goods, heavily influencing both technological trends and consumer behavior, and you want to understand the dynamics between competing models!

The two most popular smartphone models are Apple's iPhone 12 and Samsung's Galaxy S20. Each model has daily sales figures, reflecting its market demand. For example, on the first day of sales in 2020, the iPhone 12 sold 1,060 units, and the Galaxy S20 sold 991 units.

**You want to predict the sales of the Galaxy S20 on any given day based on the sales figures of the iPhone 12.** Below, we load a table called `iphone_vs_galaxy`. This table has 3 columns:
* `Date`, the date
* `iPhone Sales`,  the number of iPhone 12 units sold on that day,
* `Galaxy Sales`,  the number of Galaxy S20 units sold on that day,
* `Day`, the number of days since the start of our data

```python
iphone_vs_galaxy = Table.read_table('iPhone_vs_Galaxy_Sales.csv')
iphone_vs_galaxy.show(20)
```

<!-- BEGIN QUESTION -->

**Question 2.1.** In the cell below, create an overlaid scatter plot that visualizes the iPhones sales and Galaxy sales as a function of the day. Both iPhone and Galaxy sales should be plotted on the same graph. **(5 points)**

*Hint*: [Section 7.3](https://inferentialthinking.com/chapters/07/3/Overlaid_Graphs.html#overlaid-line-plots) in the textbook might be helpful!

*Note:* Think about how comparable iPhone 12 and Galaxy S20 sales are.


```python
# Create a scatter plot of iPhone Sales and Galaxy Sales as a function of time
...
```

<!-- END QUESTION -->

**Question 2.2.** Now, calculate the correlation coefficient between the sales of Apple iPhone 12 and Galaxy S20 using the `correlation` function you defined earlier. **(5 points)**


```python
r = ...
r
```

```python
grader.check("q2_2")
```

**Question 2.3.** Write a function `galaxy_sales_predictor` which takes daily sales figures for the iPhone 12 and predicts the daily sales of the Galaxy S20. Use the function `compute_regression_stats` that you defined earlier in this homework to assist in this process. **(5 points)**

*Hint*: Double-check what the `tbl` input to `compute_regression_stats` must look like!

*Note*: Ensure that your `galaxy_sales_predictor` function is using least squares linear regression to make predictions.



```python
def galaxy_sales_predictor(iPhone_sales):
    parameters = ...
    slope = ...
    intercept = ...
    ...
```

```python
grader.check("q2_3")
```

<!-- BEGIN QUESTION -->

**Question 2.4.** Now, using the `galaxy_sales_predictor` function you just defined, make a scatter plot with iPhone 12 sales along the x-axis and both actual and predicted Galaxy S20 sales along the y-axis. The color of the dots for the actual Galaxy S20 sales should be different from the color for the predicted Galaxy S20 sales. **(5 points)**

*Hint 1:* An example of such a scatter plot is generated can be found [here](https://inferentialthinking.com/chapters/15/2/Regression_Line.html). </a>

*Hint 2*: Think about the table that must be produced and used to generate this scatter plot. The table should include columns for iPhone 12 sales, actual Galaxy S20 sales, and predicted Galaxy S20 sales. Each row in the table should represent a day’s data. Constructing the table will be the main part of this question; once you have this table, generating the scatter plot should be straightforward as usual.


```python
iPhone_sales = ...
galaxy_pred = ...
galaxy_pred_actual = ...
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.5.** Considering the shape of the scatter plot of the true data, is the model we used reasonable? If so, what features or characteristics make this model reasonable? If not, what features or characteristics make it unreasonable? **(5 points)**


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 3. Evaluating NFL Game Predictions

#### A Brief Introduction to Sports Betting

In an NFL game, similar to basketball, the team playing at its home stadium is known as the "home team", while the team visiting the stadium is the "away team". Each team competes to score points through touchdowns, field goals, and safeties, and at the end of the game, the team with the greater number of points is declared the winner.

The difference in score can be expressed through what is known as the **outcome**, which is calculated by **subtracting the home team's score from the away team's score**:



$$\text{outcome} = \text{points scored by the away team} - \text{points scored by the home team}$$

If this number is positive, the away team won. If the number is negative, the home team won the game.

In order to facilitate betting on games, analysts at casinos try to predict the outcome of the game. This prediction of the outcome is called the **spread.** The **spread** is a forecast of the number of points by which a stronger team is expected to defeat a weaker one.

```python
nfl_spreads = Table.read_table("NFL Spread2.csv")
nfl_spreads.show(20)
```

Here's a scatter plot of the outcomes and spreads, with the spreads on the horizontal axis.

```python
nfl_spreads.scatter("Spread", "Outcome")
```

Let's investigate how well the casinos are predicting game outcomes.

One question we can ask is: Is the casino's prediction correct on average? In other words, for every value of the spread, is the average outcome of games assigned that spread equal to the spread? If not, the casino would apparently be making a systematic error in its predictions.

**Question 3.1.** Compute the correlation coefficient between outcomes and spreads. **(5 points)**

*Note:* It might be helpful to use the `correlation` function.


```python
nfl_spread_r = ...
nfl_spread_r
```

```python
grader.check("q3_1")
```

**Question 3.2.** Among games with a spread between .5 and 3.5 (including both .5 and 3.5), what was the average outcome? **(5 points)**


```python
nfl_spreads_around_2 = ...
nfl_spread_2_outcome_average = ...
print("Average outcome for spreads around 2:", nfl_spread_2_outcome_average)
```

```python
grader.check("q3_2")
```

**Question 3.3.** Using the `compute_regression_stats` function that you defined earlier, compute the least squares linear regression line for predicting the actual game outcomes from the provided betting spreads. We have prepared a dataset for you, structured in a two-column format; the first column is labeled `Spread` (predictor variable, x), and the second column is labeled `Outcome` (response variable, y). Use this dataset as an argument to your function to carry out the regression analysis. **(5 points)**



```python
evaluate_table = nfl_spreads.select('Spread', 'Outcome')
approximate = ...
nfl_spread_slope = ...
nfl_spread_intercept = ...
print("Slope:", round(nfl_spread_slope, 3))
print("Intercept", round(nfl_spread_intercept, 3))
```

```python
grader.check("q3_3")
```

<!-- BEGIN QUESTION -->

**Question 3.4.** 
Do you think this regression line will provide casino analysts with a reliable prediction of game outcomes? Provide a rationale for your answer. Include a discussion of the correlation coefficient and the scatter plot in your answer. **(5 points)**


_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Fitting a Least-Squares Regression Line

Recall that the least-squares regression line is the unique straight line that minimizes root mean squared error (RMSE) among all possible fit lines. Using this property, we can find the equation of the regression line by finding the pair of slope and intercept values that minimize root mean squared error. 

**Question 3.5.** Define a function called `errors`.  It should take three arguments:
1. a table `table` like `spreads` (with the same column names and meanings, but not necessarily the same data)
2. the `slope` of a line (a number)
3. the `intercept` of a line (a number).

It should **return an array of the errors** made when a line with that slope and intercept is used to predict outcome from spread for each game in the given table. **(3 points)**

*Note*: Make sure you are returning an array of the errors, and not the RMSE. 


```python
def errors(table, slope, intercept):
    ...
```

```python
grader.check("q3_5")
```

**Question 3.6.** Using `errors`, compute the errors for the line with slope `2` and intercept `35` on the `spreads` dataset. Name that array `outcome_errors`.  Then, make a scatter plot of the errors. **(3 points)**

*Hint:* To make a scatter plot of the errors, plot the error for each outcome in the dataset.  Put the actual spread on the horizontal axis and the outcome error on the vertical axis.


```python
outcome_errors = ...
...
```

```python
grader.check("q3_6")
```

You should find that the errors are almost all negative.  That means our line is not the best fit to our data.  Let's find a better one.

**Question 3.7.** Define a function called `fit_line`.  It should take a table like `nfl_spreads` (with the same column names and meanings) as its argument.  It should return an array containing the slope (as the first element) and intercept (as the second element) of the least-squares regression line predicting outcome from spread for that table. **(3 points)**

*Hint*: Define a function `rmse` within `fit_line` that takes a slope and intercept as its arguments. `rmse` will use the table passed into `fit_line` to compute predicted outcomes and then return the root mean squared error between the predicted and actual outcomes. Within `fit_line`, you can call `rmse` the way you would any other function.

If you haven't tried to use the `minimize` [function](http://data8.org/sp22/python-reference.html) yet, now is a great time to practice. Here's an [example from the textbook](https://www.inferentialthinking.com/chapters/15/3/Method_of_Least_Squares.html#numerical-optimization).


```python
def fit_line(tbl):
    # Your code may need more than 1 line below here.
    def rmse(..., ...):
        return ... 
    return ... 
    
# Here is an example call to your function.  To test your function,
# figure out the right slope and intercept by hand.
example_table = Table().with_columns(
    "Spread", make_array(0, 1),
    "Outcome", make_array(1, 3))
fit_line(example_table)
```

```python
grader.check("q3_7")
```

**Question 3.8.** Use `fit_line` to fit a line to `spreads`, and assign the output to `regression_line`. Assign the first and second elements in `regression_line` to `regression_line_slope` and `regression_line_intercept`, respectively.

Then, set `new_errors` to the array of errors that we get by calling `errors` with our new regression line. The provided code will graph the corresponding residual plot with a best fit line. **(3 points)**

*Hint:* Make sure that the residual plot makes sense. What qualities should the best fit line of a residual plot have?


```python
regression_line = ...
regression_line_slope = ...
regression_line_intercept = ...

new_errors = ...

# This code displays the residual plot, given your values for the regression_line_slope and regression_line_intercept
Table().with_columns("Spread", 
                    nfl_spreads.column("Spread"), 
                    "Outcome errors", 
                    new_errors
                   ).scatter("Spread", "Outcome errors", fit_line=True)

# This just prints your slope and intercept
"Slope: {:g} | Intercept: {:g}".format(regression_line_slope, regression_line_intercept)
```

```python
grader.check("q3_8")
```

<!-- BEGIN QUESTION -->

**Question 3.9.** How does the process of minimizing the Root Mean Square Error (RMSE) lead to slope and intercept values in Question 3.8 that are consistent with those obtained using previous methods in Question 3.3, even though different techniques are employed? **(3 points)**


_Type your answer here, replacing this text._

<!-- END QUESTION -->

You're done with Homework 10!  

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [0, 0, 0])
```
>>> std_units([1,2,3,4,5])
array([-1.41421356, -0.70710678,  0.        ,  0.70710678,  1.41421356])
```
```
>>> np.isclose(np.mean(std_units([1,2,3,4,5])), 0)
True
```
```
>>> np.isclose(np.std(std_units([1,2,3,4,5])), 0.9999999999999999)
True
```


Tests for q1_2 (points: [0])
```
>>> np.isclose(correlation([1,2,3], [4,5,6]), 0.9999999999999999)
True
```


Tests for q1_5 (points: [0, 0, 0])
```
>>> len(constraints) == 3
True
```
```
>>> # Make sure your function is returning an array!
>>> type(constraints) == np.ndarray
True
```
```
>>> np.isclose(constraints.item(0), 0.11003088267167993)
True
```


Tests for q1_6 (points: [0, 0])
```
>>> # Make sure regression_changes is a numpy array with 3 items!
>>> len(regression_change) == 2 and type(regression_change) == np.ndarray
True
```
```
>>> regression_change.item(0) in set([0, 1, False, True])
True
```


Tests for q1_7 (points: [0])
```
>>> 35 <= free_throw_3pt_est <= 45
True
```


Tests for q2_2 (points: [0, 0])
```
>>> # Correlation is a number between -1 and 1
>>> -1 <= r <= 1
True
```
```
>>> # It appears that you implemented std_units, but did so incorrectly
>>> std_units(np.arange(5)) is None or np.allclose(std_units(np.arange(5)), [-1.41421356, -0.70710678,  0,  0.70710678,  1.41421356])
True
```


Tests for q2_3 (points: [0, 2])
```
>>> galaxy_sales_predictor(np.arange(10,15))
array([ 308.08464403,  308.73092002,  309.37719601,  310.02347201,
        310.669748  ])
```
```
>>> 0 < galaxy_sales_predictor(10) < 400
True
```


Tests for q3_1 (points: [0])
```
>>> -1 <= nfl_spread_r <= 1
True
```


Tests for q3_2 (points: [0])
```
>>> -4 <= nfl_spread_2_outcome_average <= nfl_spread_2_outcome_average
True
```


Tests for q3_3 (points: [0])
```
>>> -1 <= nfl_spread_slope < .6
True
```


Tests for q3_5 (points: [0, 0])
```
>>> type(errors(nfl_spreads, 0, 0).item(0)) in set([float, np.float32, np.float64])
True
```
```
>>> len(errors(nfl_spreads, 0, 0)) == 1083
True
```


Tests for q3_6 (points: [3])
```
>>> # If you're not passing this test, double check your errors function!
>>> np.round(outcome_errors.item(0), 2) == -27.0
True
```


Tests for q3_7 (points: [0, 0])
```
>>> type(fit_line(example_table)) == np.ndarray
True
```
```
>>> type(fit_line(example_table).item(0)) in set([float, np.float32, np.float64])
True
```


Tests for q3_8 (points: [0])
```
>>> len(new_errors) == 1083
True
```
