```python
# Initialize Otter
import otter
grader = otter.Notebook("hw07.ipynb")
```

# Homework 7: Testing Hypotheses

Please complete this notebook by filling in the cells provided. Before you begin, execute the previous cell to load the provided tests.

**Helpful Resource:**

- [Python Reference](http://data8.org/sp22/python-reference.html): Cheat sheet of helpful array & table methods used in Data 8!

**Recommended Readings**: 

* [Sampling Methods Guide](https://drive.google.com/file/d/1UtNdxDdI5XphWvwgTjODiAUFwArpZtoK/view)
* [Testing Hypotheses](https://www.inferentialthinking.com/chapters/11/Testing_Hypotheses.html)
* [A/B Testing](https://inferentialthinking.com/chapters/12/1/AB_Testing.html)

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute this cell again.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. **Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Note: This homework has hidden tests on it. That means even though the tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**


Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

You should start early so that you have time to get help if you're stuck.

```python
# Run this cell to set up the notebook, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *
import d8error

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

## 1. Coffee Drinkers Delight

A popular coffee shop introduces two new seasonal drinks: Drink A and Drink B. Customers visiting the shop are offered one of the two drinks based on the day of the week. Those who visit on Thursday, Friday, Saturday, and Sunday are served Drink A, while visitors on Monday, Tuesday and Wednessday receive Drink B. The coffee shop is never closed.

Barista Alex remarks, "It's practically like flipping a coin where it lands on heads $\frac{4}{7}$ of the time. Land on heads, you get Drink A; tails, you get Drink B."

However, Barista Jordan disagrees, stating, "It's not similar to flipping a coin at all. Our approach is entirely different."

During one particular week, the coffee shop served Drink A to 260 customers and Drink B to 140 customers. Your task is to conduct a hypothesis test to determine which barista's viewpoint is more consistent with the observed data.



**Question 1.1.** Given the information above, what was the sample size for the data, and what was the percentage of people who got **Drink A?** **(4 points)**

*Note*: Your percent should be a number between 0 and 100.


```python
sample_size = ...
percent_DA = ...

print(f"Sample Size: {sample_size}")
print(f"Drink A Percent: {percent_DA}")
```

```python
grader.check("q1_1")
```

<!-- BEGIN QUESTION -->

**Question 1.2.** State the null hypothesis. It should reflect the position of either Barista Alex or Barista Jordan. **(4 points)**

*Note:* Check out [11.3](https://inferentialthinking.com/chapters/11/3/Decisions_and_Uncertainty.html#step-1-the-hypotheses) for a refresher on hypotheses.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.3.** State the alternative hypothesis. It should reflect the position of the Barista you did not choose to represent in Question 1.2. **(4 points)**

*Note:* Check out [11.3](https://inferentialthinking.com/chapters/11/3/Decisions_and_Uncertainty.html#step-1-the-hypotheses) for a refresher on hypotheses.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 1.4.** One of the test statistics below is appropriate for testing these hypotheses. Assign the variable `test_stat` to the number corresponding to the correct test statistic. **(4 points)**

*Note:* For simplicity purposes, all values are rounded to the nearest integer. 

1. |percent of heads - 57|
2. |percent of heads - 50|
3. percent of heads - 57
4. percent of heads - 50


```python
test_stat = ...
test_stat
```

```python
grader.check("q1_4")
```

**Question 1.5.** Using your answer from Questions 1.1 and 1.4, find the observed value of the test statistic and assign it to the variable `observed_stat`. **(4 points)**


```python
observed_stat = ...
observed_stat
```

```python
grader.check("q1_5")
```

**Question 1.6.** In order to perform this hypothesis test, you must simulate the test statistic. From the four options below, pick the assumption that is needed for this simulation. Assign `assumption_required` to an integer corresponding to the assumption. **(4 points)**

1. No assumptions are needed. We can just simulate the statistic.
2. The statistic must be simulated under both hypotheses.
3. The statistic must be simulated under the alternative hypothesis.
4. The statistic must be simulated under the null hypothesis.


```python
assumption_required = ...
assumption_required
```

```python
grader.check("q1_6")
```

<!-- BEGIN QUESTION -->

**Question 1.7.** Simulate 15,000 values of the test statistic under the assumption you picked in Question 1.6. **(4 points)** 

As usual, start by defining a function that simulates one value of the statistic. Your function should use `sample_proportions`. (You may find a variable defined in Question 1.1 useful here!) Then, write a `for` loop to simulate multiple values and collect them in the array `simulated_statistics`.

Use as many lines of code as you need. We have included the code that visualizes the distribution of the simulated values. The red dot represents the observed statistic you found in Question 1.5.


```python
def one_simulated_statistic():
    ...

num_simulations = 15000

simulated_statistics = ...
for ... in ...:
    ...

# Run the this cell a few times to see how the simulated statistic changes
one_simulated_statistic()
```

<!-- END QUESTION -->



```python
# Run this cell to produce a histogram of the simulated statistics

Table().with_columns('Simulated Statistic', simulated_statistics).hist()
plt.scatter(observed_stat, -0.002, color='red', s=40);
```

**Question 1.8.** Using `simulated_statistics`, `observed_stat`, and `num_simulations`, find the empirical p-value based on the simulation. **(4 points)**


```python
p_value = ...
p_value
```

```python
grader.check("q1_8")
```

**Question 1.9.** Assign `correct_barista` to the number corresponding to the correct statement below. Use the 5% cutoff for the p-value. **(4 points)**

1. The data support Barista Alex's position more than they support Barista Jordan's.
2. The data support Barista Jordan's position more than they support Barista Alex's.

As a reminder, here are the two claims made by Barista Alex and Barista Jordan:
> **barista Alex:** "It's practically like flipping a coin where it lands on heads $\frac{4}{7}$ of the time. Land on heads, you get Drink A; tails, you get Drink B."

>**Barista Jordan:** "It's not similar to flipping a coin at all. Our approach is entirely different."


```python
correct_doctor = ...
correct_doctor
```

```python
grader.check("q1_9")
```

## 2. Using TVD as a Test Statistic

Before beginning this section, please read [this section](https://inferentialthinking.com/chapters/11/2/Multiple_Categories.html#a-new-statistic-the-distance-between-two-distributions) of the textbook on TVD!

**Total variation distance (TVD)** is a special type of test statistic that we use when we want to compare two distributions of *categorical data*. It is often used when we observe that a set of observed proportions/probabilities is different than what we expect under the null model. 

Consider a six-sided die that we roll 6,000 times. If the die is fair, we would expect that each face comes up $\frac{1}{6}$ of the time. By random chance, a fair die won't always result in equal proportions (that is, we won't get exactly 1,000 of each face). However, if we suspect that the die might be unfair based on the data, we can conduct a hypothesis test using TVD to compare the expected [$\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$, $\frac{1}{6}$] distribution to what is actually observed.

In this part of the homework, we'll look at how we can use TVD to determine the effect that different factors have on happiness. 

We will be working with data from the [Gallup World Poll](https://www.gallup.com/analytics/349487/gallup-global-happiness-center.aspx#:~:text=World%20Happiness%20Report&text=Using%20the%20Gallup%20World%20Poll,about%20the%20World%20Happiness%20Report.) that is presented in the World Happiness Report, a survey of the state of global happiness. The survey ranked 155 countries by overall happiness and estimated the influence that economic production, social support, life expectancy, freedom, absence of corruption, and generosity had on population happiness. The study has been repeated for several years, but we'll be looking at data from the 2016 survey.

Run the cell below to load in the `happiness_scores` table.

```python
happiness_scores = Table.read_table("happiness_scores.csv")
happiness_scores.show(5)
```

Participants in the study were asked to evaluate their life satisfaction from a scale of 0 (worst possible life) to 10 (best possible life). The responses for each country were averaged to create the `Happiness Score`.

The columns `Economy (GDP per Capita)`, `Family`, `Health (Life Expectancy)`, `Freedom`, `Trust (Government Corruption)`, and `Generosity` estimate the extent to which each factor influences happiness, both for better or for worse. The happiness score is the sum of these factors; the larger a factor is, the more it contributes to overall happiness. [In other words, if you add up all the factors (in addition to a "Difference from Dystopia" value we excluded in the dataset), you get the happiness score.]

Let's look at the different factors that affect happiness in the United States. Run the cell below to view the row in `us_happiness` that contains data for the United States.

```python
us_happiness = happiness_scores.where("Country", "United States")
us_happiness
```

**To compare the different factors, we'll look at the proportion of the happiness score that is attributed to each variable. 
You can find these proportions in the table `us_happiness_factors` after running the cell below.**

*Note:* The factors shown in `us_happiness` don't add up exactly to the happiness score, so we adjusted the proportions to  only account for the data we have access to. The proportions were found by dividing each Happiness Factor value by the sum of all Happiness Factor values in `us_happiness`.

```python
us_happiness_factors = Table().read_table("us_happiness_factors.csv")
us_happiness_factors
```

<!-- BEGIN QUESTION -->

**Question 2.1.** Suppose we want to test whether or not each factor contributes the same amount to the overall Happiness Score. Define the null hypothesis, alternative hypothesis, and test statistic in the cell below. Feel free to check your work with another student or course staff. **(4 points)**

*Note:* Please format your answer as follows:
- Null Hypothesis: ...  
- Alternative Hypothesis: ...  
- Test Statistic: ...  


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 2.2.** Write a function `calculate_tvd` that takes in the observed distribution (`obs_dist`) and expected distribution under the null hypothesis (`null_dist`) and calculates the total variation distance. Use this function to set `observed_tvd` to be equal to the observed test statistic. **(4 points)**


```python
null_distribution = np.ones(6) * (1/6)

def calculate_tvd(obs_dist, null_dist):
    ...
    
observed_tvd = ...
observed_tvd
```

```python
grader.check("q2_2")
```

**Question 2.3.** Create an array called `simulated_tvds` that contains 10,000 simulated values under the null hypothesis. Assume that the original sample consisted of 1,000 individuals. **(4 points)**

*Hint:* The `sample_proportions` function may be helpful to you. Refer to the [Python Reference Sheet](http://data8.org/sp22/python-reference.html#:~:text=sample_proportions(sample_size%2C%20model_proportions)) to read up on it!


```python
simulated_tvds = ...

...
```

```python
grader.check("q2_3")
```

Run the cell below to plot a histogram of your simulated test statistics, as well as a red dot representing the observed value of the test statistic.

```python
Table().with_column("Simulated TVDs", simulated_tvds).hist()
plt.scatter(observed_tvd, 0, color='red', s=70, zorder=2);
plt.show();
```

**Question 2.4.** Use your simulated statistics to calculate the p-value of your test. Make sure that this number is consistent with what you observed in the histogram above. **(4 points)**


```python
p_value_tvd = ...
p_value_tvd
```

```python
grader.check("q2_4")
```

<!-- BEGIN QUESTION -->

**Question 2.5.** What can you conclude about how each factor contributes to the overall happiness score in the US? Explain your answer using the results of your hypothesis test. Assume a p-value cutoff of 5%. **(4 points)**


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 3. Fitness Level Comparison

A group of fitness enthusiasts conducted a survey to compare physical fitness levels among individuals. They managed to gather a simple random sample of 300 adults from a broad demographic. Each participant identified as either "active" or "sedentary" based on their lifestyle choices. Information was collected on various aspects, including the participants' fitness scores. The dataset, named `fitness_scores`, comprises one row per participant, with columns detailing each individual's lifestyle choice.



```python
lifestyle_fitness_table = Table.read_table('lifestyles.csv')
lifestyle_fitness_table.show(5)
```

**Question 3.1.** How many people were sedentary in our sample? Please use the provided skeleton code. **(4 points)**

*Hint:* Keep in mind that `.group` sorts categories in alphabetical order!


```python
num_sedentary = lifestyle_fitness_table.group(...)...
num_sedentary
```

```python
grader.check("q3_1")
```

**Question 3.2.** Complete the cell below so that `avg_active_vs_sedentary` evaluates to `True` if the sampled active people have an higher fitness score than the sampled sedentary on average, and `False` otherwise. Use Python code to achieve this. **(4 points)**


```python
group_mean_tbl = lifestyle_fitness_table.group(...)
group_means = group_mean_tbl...       # array of mean ages
avg_active_vs_sedentary = group_means... > group_means...
avg_active_vs_sedentary
```

```python
grader.check("q3_2")
```

**Question 3.3.** The researchers want to use the data to test whether individuals with an "active" lifestyle have higher fitness scores than those with a "sedentary" lifestyle—or, in other words, whether the fitness scores of the two groups have the same distribution. One of the following statements is their null hypothesis and another is their alternative hypothesis. Assign `null_statement_number` and `alternative_statement_number` to the numbers of the correct statements in the code cell below. **(4 points)**


1. In the population, "active" and "sedentary" individuals have the same distribution of fitness scores; any difference in sample averages of the two groups is due to chance.
2. In the sample, "active" and "sedentary" individuals have the same distribution of fitness scores; any difference in sample averages of the two groups is due to chance.
3. The fitness score distributions of "active" and "sedentary" individuals in the population are different due to chance.
4. "Active" individuals in the population have higher fitness scores than "sedentary" individuals, on average.
5. "Active" individuals in the sample have higher fitness scores than "sedentary" individuals, on average.
6. The average fitness scores of "active" and "sedentary" individuals in the population are different.




```python
null_statement_number = ...
alternative_statement_number = ...
```

```python
grader.check("q3_3")
```

**Question 3.4.** The researchers have decided to use a permutation test for their analysis. Assign `permutation_test_reason` to the number corresponding to the reason they made this choice. **(4 points)**


1. Since a person's fitness score shouldn't be inherently related to their lifestyle category, it doesn't matter who is labeled "active" and who is labeled "sedentary", so you can use permutations.
2. Under the null hypothesis, permuting the rows of `lifestyle_fitness_table` is equivalent to drawing a new random sample with the same number of "active" and "sedentary" individuals as in the original sample.
3. Under the null hypothesis, permuting the labels in the `lifestyle_fitness_table` is equivalent to drawing a new random sample with the same number of "active" and "sedentary" individuals as in the original sample.


```python
permutation_test_reason = ...
permutation_test_reason
```

```python
grader.check("q3_4")
```

**Question 3.5.** To test their hypotheses, the researchers have followed our textbook's advice and chosen a test statistic where the following statement is true: Large values of the test statistic favor the alternative hypothesis.


The researchers' test statistic is one of the two options below. Which one is it? Assign the appropriate number to the variable `correct_test_stat.` **(4 points)**


1. "|active fitness score average - sedentary fitness score average|" in a sample created by randomly shuffling the "active"/"sedentary" labels.
2. "active fitness score average - sedentary fitness score average" in a sample created by randomly shuffling the "active"/"sedentary" labels.


```python
correct_test_stat = ...
correct_test_stat
```

```python
grader.check("q3_5")
```

**Question 3.6.** Complete the cell below so that `observed_statistic_fitness` evaluates to the observed value of the researchers' test statistic based on the fitness scores between "active" and "sedentary" lifestyle groups. Use as many lines of code as you need, and remember that you can use any quantity, table, or array that you created earlier. **(4 points)**

```python
observed_statistic_fitness = ...
observed_statistic_fitness
```

```python
grader.check("q3_6")
```

**Question 3.7.** Assign `shuffled_labels` to an array of shuffled active/sedentary labels. The rest of the code puts the array in a table along with the data in `lifestyle_fitness_table`. **(4 points)**

*Note:* Check out [12.1](https://inferentialthinking.com/chapters/12/1/AB_Testing.html#predicting-the-statistic-under-the-null-hypothesis) for a refresher on random permutations.


```python
shuffled_labels = ...
original_with_shuffled_labels = lifestyle_fitness_table.with_columns('Shuffled Label', shuffled_labels)
original_with_shuffled_labels
```

```python
grader.check("q3_7")
```

**Question 3.8.** The comparison below uses the array `shuffled_labels` from Question 3.7 and the count `num_sedentary` from Question 3.1.

For this comparison, assign the correct number from one of the following options to the variable `correct_q8`. **Pretend this is a midterm problem and solve it without doing the calculation in a code cell.** **(4 points)**

`comp = np.count_nonzero(shuffled_labels == 'sedentary') == num_sedentary`

1. `comp` is set to `False`.
2. `comp` is set to `True`.
3. `comp` is set to `True` or `False`, depending on how the shuffle came out.


```python
correct_q8 = ...
correct_q8
```

```python
grader.check("q3_8")
```

**Question 3.9.** Create a function named `generate_one_simulation` that does not accept any arguments and returns a single simulated value of the test statistic based on the comparison of fitness scores between "active" and "sedentary" lifestyle categories. We have outlined a basic structure for you below, but you are encouraged to adapt this framework in a way that best suits your analysis. Feel free to utilize as many lines of code as necessary. Consider revisiting your previous code from this assignment, as there may be elements that you can efficiently incorporate into this function. **(4 points)**




```python
def generate_one_simulation():
    "Returns one value of our simulated test statistic"
    shuffled_labels = ...
    shuffled_tbl = ...
    group_means = ...
    ...
```

```python
grader.check("q3_9")
```

After you have defined your function, run the following cell a few times to see how the statistic varies.

```python
 generate_one_simulation()
```

**Question 3.10.** Write the code to generate 5,000 simulated values for the test statistic, which compares the fitness scores between the "active" and "sedentary" lifestyle groups. Included in this task is the code to plot the empirical distribution of these simulated statistics alongside the `observed_statistic_fitness` value you calculated in Question 3.6. You are encouraged to utilize as many lines of code as necessary for this simulation. **(4 points)**

*Note:* Expect this code to take approximately one minute to execute.

```python
repetitions = 5000

generated_simulations_ab = make_array()
...
    generated_simulations_ab = ...

# Do not change these lines
Table().with_columns('Generated Simulation',generated_simulations_ab ).hist()
plt.scatter(observed_statistic_fitness, -0.002, color='red', s=70);
```

```python
grader.check("q3_10")
```

**Question 3.11.** Use the simulation to find an empirical approximation to the p-value. Assign `p_val` to the appropriate p-value from this simulation. Then, assign `conclusion` to either `null_hyp` or `alt_hyp`. **(4 points)** 

*Note:* Assume that we use the 5% cutoff for the p-value.


```python
# These are variables provided for you to use.
null_hyp = 'The data are consistent with the null hypothesis.'
alt_hyp = 'The data support the alternative more than the null.'

p_val = ...
conclusion = ...

p_val, conclusion # Do not change this line
```

```python
grader.check("q3_11")
```

You're done with Homework 7!  

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [0, 0, 0, 0])
```
>>> type(sample_size) == int
True
```
```
>>> type(percent_DA) == float or type(percent_DA) == int
True
```
```
>>> 100 <= sample_size <= 500
True
```
```
>>> 0 <= percent_DA <= 100
True
```


Tests for q1_4 (points: [0, 0])
```
>>> type(test_stat) == int
True
```
```
>>> any((test_stat == x for x in np.arange(1,5)))
True
```


Tests for q1_5 (points: [0, 0])
```
>>> type(observed_stat) == float
True
```
```
>>> 0 <= observed_stat <= 100
True
```


Tests for q1_6 (points: [0, 0])
```
>>> type(assumption_required) == int
True
```
```
>>> 1 <= assumption_required <= 4
True
```


Tests for q1_8 (points: [0, 0])
```
>>> import numpy as np
>>> type(p_value) == float or type(p_value) == np.float64
True
```
```
>>> 0 <= p_value <= 1
True
```


Tests for q1_9 (points: [0, 0])
```
>>> type(correct_doctor) == int
True
```
```
>>> any((correct_doctor == x for x in (1,2)))
True
```


Tests for q2_2 (points: [2, 2])
```
>>> import numpy as np
>>> type(observed_tvd) in set([float, np.float32, np.float64])
True
```
```
>>> np.isclose(round(observed_tvd, 6), 0.262086)
True
```


Tests for q2_3 (points: [1, 1, 1, 1])
```
>>> len(simulated_tvds) == 10000
True
```
```
>>> all(simulated_tvds >= 0)
True
```
```
>>> # Make sure that your test statistic is not always the same value 
>>> len(np.unique(simulated_tvds)) != 1
True
```
```
>>> # Make sure you're simulating under the null hypothesis
>>> np.mean(simulated_tvds) < 0.1
True
```


Tests for q2_4 (points: [4])
```
>>> 0 <= p_value_tvd <= 1
True
```


Tests for q3_1 (points: [0])
```
>>> 0 <= num_sedentary <= 300
True
```


Tests for q3_10 (points: [4])
```
>>> len(generated_simulations_ab) == 5000
True
```


Tests for q3_11 (points: [0, 0, 0])
```
>>> import numpy as np
>>> type(p_val) in set([float, np.float32, np.float64])
True
```
```
>>> 0 <= p_val <= 1
True
```
```
>>> type(conclusion) == str
True
```


Tests for q3_2 (points: [0])
```
>>> import numpy as np
>>> type(avg_active_vs_sedentary) in set([bool, np.bool_])
True
```


Tests for q3_3 (points: [0, 0, 0, 0, 0])
```
>>> type(null_statement_number) == int
True
```
```
>>> type(alternative_statement_number) == int
True
```
```
>>> any(null_statement_number == x for x in np.arange(1,7))
True
```
```
>>> any(alternative_statement_number == x for x in np.arange(1,7))
True
```
```
>>> null_statement_number != alternative_statement_number
True
```


Tests for q3_4 (points: [0, 0])
```
>>> type(permutation_test_reason) == int
True
```
```
>>> any(permutation_test_reason == x for x in np.arange(1,4))
True
```


Tests for q3_5 (points: [0, 0])
```
>>> type(correct_test_stat) == int
True
```
```
>>> any(correct_test_stat == x for x in np.arange(1,3))
True
```


Tests for q3_6 (points: [0, 0])
```
>>> type(observed_statistic_fitness) == float
True
```
```
>>> observed_statistic_fitness >= 0
True
```


Tests for q3_7 (points: [0, 0, 0])
```
>>> type(original_with_shuffled_labels) == Table
True
```
```
>>> original_with_shuffled_labels.labels == ("Lifestyle", "Fitness Score", "Shuffled Label")
True
```
```
>>> original_with_shuffled_labels.num_rows == 300
True
```


Tests for q3_8 (points: [0, 0])
```
>>> type(correct_q8) == int
True
```
```
>>> any(correct_q8 == x for x in np.arange(1,4))
True
```


Tests for q3_9 (points: [4])
```
>>> -5 < generate_one_simulation() < 5
True
```
