```python
# Initialize Otter
import otter
grader = otter.Notebook("lab08-sp24.ipynb")
```

# Lab 8: Normal Distribution and Variability of Sample Means

Welcome to Lab 8!

In today's lab, we will learn about [the variability of sample means](https://inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html) as well as [the normal distribution](https://www.inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html).


**Submission**: Once you’re finished, run all cells besides the last one, select File > Save Notebook, and then execute the final cell. Then submit the downloaded zip file, that includes your notebook,  according to your instructor's directions. 

First, set up the notebook by running the cell below.

```python
# Run this cell, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore')

import d8error
```

# 1. Normal Distributions

When we visualize the distribution of a sample, we are often interested in the mean and the standard deviation of the sample (for the rest of this lab, we will abbreviate “standard deviation” as “SD”). These two sample statistics can give us a bird’s eye view of the distribution -- by letting us know where the distribution sits on the number line and how spread out it is, respectively. 

We want to check if the data is linearly related, so we should look at the data.

**Question 1.1.** FIFA 23 is a soccer simulation video game developed and published by Electronic Arts under the EA Sports label. The next cell loads a table called `FIFA23` which contains over 17,000 unique players and information about each one, such as the player's age, nationality, and overall rating.

Plot the distribution of player’s ages from the table. Don’t change the last line, which will plot the mean of the sample on the distribution itself.


```python
FIFA = Table.read_table('FIFA23_official_data.csv')
...

# Do not change anything below this line
plots.scatter(np.mean(FIFA.column("Age")), 0, color='red', zorder=2, s=50, label="Mean");
plots.ylim(-0.005, 0.07);
```

From the plot above, we can see that the mean is the center of gravity or balance point of the distribution. If you cut the distribution out of cardboard, and then placed your finger at the mean, the distribution would perfectly balance on your finger. Since the distribution above is right skewed (which means it has a long right tail), we know that the few large values pull the mean higher. Thus, the mean of the distribution is larger than the median, which is the “halfway” point of the data. Conversely, if the distribution had been left skewed, the few small values would pull the mean lower than the median.

**Question 1.2.** Run the following cell to compare the mean (red) and median (blue) of the distribution of players ages.


```python
FIFA.hist("Age")
plots.scatter(np.mean(FIFA.column("Age")), 0, color='red', zorder=2, s=50, label="Mean")
plots.scatter(np.median(FIFA.column("Age")), 0, color='blue', zorder=2, s=50, label="Median")
plots.legend()
plots.ylim(-0.005, 0.07);
```

We are also interested in the standard deviation of player’s ages. The SD gives us a sense of how variable players' ages are around the average players' age. If the SD is large, then the players' ages should spread over a large range from the mean. If the SD is small, then the players' ages should be tightly clustered around the average player age. 

**The SD of an array is defined as the root mean square of the deviations where the deviations are the difference between the original elements and their average.**

You can read more about Variability and SD in [Chapter 14.2](https://inferentialthinking.com/chapters/14/2/Variability.html). Fun fact! $\sigma$ (Greek letter sigma) is used to represent the SD and $\mu$ (Greek letter mu) is used for the mean. 

**Question 1.3.** Assign `mean_age` to the mean of the players ages and `sd_age` to the SD of the players ages. Then run the cell below to see the width of one SD (blue) from the sample mean (red) plotted on the histogram of players ages.


```python
mean_age = ...
sd_age = ...
FIFA.hist("Age")


# Do not change anything below this line
plots.scatter(mean_age, 0, color='red', zorder=2, s=50);
plots.scatter(mean_age+sd_age, 0, marker='^', color='blue', zorder=2, s=50);
plots.scatter(mean_age-sd_age, 0, marker='^', color='blue', zorder=2, s=50);
plots.ylim(-0.005, 0.07);
```

```python
grader.check("q1_3")
```

In this histogram, the standard deviation is not easy to identify just by looking at the graph.

However, the distributions of some variables allow us to easily spot the standard deviation on the plot. For example, if a sample follows a *normal distribution*, the standard deviation is easily spotted at the point of inflection (the point where the curve begins to change the direction of its curvature) of the distribution. 

**Question 1.4.** Fill in the following code to examine the distribution of overall player ratings, which is roughly normally distributed. We’ll plot the standard deviation on the histogram, as before - notice where one standard deviation (blue) away from the mean (red) falls on the plot.


```python
overall_mean = ...
overall_sd = ...
FIFA.hist("Overall", bins=np.arange(45,85,2))

plots.scatter((overall_mean), 0, color='red', zorder=2, s=50);
plots.scatter(overall_mean+overall_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.scatter(overall_mean-overall_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.ylim(-0.01, 0.06);
```

```python
grader.check("q1_4")
```

We don’t always know how a variable will be distributed, and making assumptions about whether or not a variable will follow a normal distribution is dangerous. However, the **Central Limit Theorem** defines one distribution that always follows a normal distribution. **The distribution of the *sums* and *means* of all large random samples drawn with replacement from a single distribution (regardless of the distribution’s original shape) will be normally distributed.** Remember that the Central Limit Theorem refers to the distribution of a *statistic* calculated from a distribution, not the distribution of the original sample or population. If this is confusing, ask a staff member!

The next section will explore distributions of sample means, and you will see how the standard deviation of these distributions depends on sample sizes.

# 2. Variability of the Sample Mean

By the [Central Limit Theorem](https://www.inferentialthinking.com/chapters/14/4/Central_Limit_Theorem.html), the probability distribution of the mean or sum of a large random sample is roughly normal. The bell curve is centered at the population mean. Some of the sample means are higher and some are lower, but the deviations from the population mean are roughly symmetric on either side, as we have seen repeatedly. 

In our simulations, we also noticed that the means of larger samples tend to be more tightly clustered around the population mean than means of smaller samples. In this section, we will quantify the [variability of the sample mean](https://www.inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html) and develop a relation between the variability and the sample size.

Let's take a look at the payroll for employees of Manhattan, New York City in 2016. The mean salary of the employees reported by the city government was about $65,890.00.

*Note: If you get stuck on any part of this lab, please refer to [Chapter 14 of the textbook](https://www.inferentialthinking.com/chapters/14/Why_the_Mean_Matters.html).*

```python
NYC_payroll = Table.read_table('NYC_payroll_2016.csv').select("Base Salary")
NYC_payroll.show(5)
```

```python
NYC_payroll_mean = np.mean(NYC_payroll.column('Base Salary'))
NYC_payroll_mean
print('Mean salary of Manhattan employees in 2016: ', NYC_payroll_mean.round(2))
```

```python
NYC_payroll.hist('Base Salary', bins=np.arange(0, 220000+10000*2, 10000))
plots.scatter(NYC_payroll_mean, 0, color='red', zorder=2, s=50)
plots.title('2016 salaries of NYC employees')
plots.ylim(-0.0000005, 0.000025);
```

Clearly, the population *does not* follow a normal distribution. Keep that in mind as we progress through these exercises.

Let's take random samples **with replacement** and look at the probability distribution of the sample mean. As usual, we will use simulation to get an empirical approximation to this distribution.

**Question 2.1.** Define a function `get_sample_mean`. Its arguments should be `table` (the name of a table), `label` (the label of the column containing the variable), and `sample_size`(the number of employees in the sample). It should sample with replacement from the table and
return the mean of the `label` column of the sample.


```python
def get_sample_mean(table, label, sample_size):
    sample_data = ...
    sample_data_mean = ...
    ...
```

```python
grader.check("q2_1")
```

**Question 2.2.** Use `get_sample_mean` to define a function `simulate_sample_mean`. The arguments are the name of the `table`, the `label` of the column containing the variable, the `sample_size`, and the number of simulations (`repetitions`).

The function should sample with replacement from the table and calculate the mean of each sample. It should save the sample means in an array called `means_array`. The remaining code in the function displays an empirical histogram of the sample means.

*Note: We round all values to 2 decimal places because we are working with salaries.*


```python
"""Empirical distribution of random sample means"""

def simulate_sample_mean(table, label, sample_size, repetitions):
    
    means_array = make_array()

    for i in np.arange(repetitions):
        new_sample_mean = ...
        means_array = ...

    means_table = Table().with_column('Sample Means', means_array)
    
    # Display empirical histogram and print all relevant quantities – don't change this!
    means_table.hist(bins=20)
    plots.xlabel('Sample Means')
    plots.title('Sample Size {sample_size}; {repetitions} Resamples'.format(sample_size=sample_size, repetitions=repetitions))
    print("Sample size: ", sample_size)
    print("Population mean:", np.mean(table.column(label)).round(2))
    print("Average of sample means: ", np.mean(means_array).round(2))
    print("Population SD:", np.std(table.column(label)).round(2))
    print("SD of sample means:", np.std(means_array).round(2))
    return np.std(means_array)
```

Verify with a peer or TA that you've implemented the function above correctly. If you haven't implemented it correctly, the rest of the lab won't work properly, so this step is crucial.

 In the following cell, we will create a sample of size 100 from `salaries` and graph it using our new `simulate_sample_mean` function.
 
 *Hint: You should see a distribution similar to something we've been talking about. If not, check your function.*

```python
simulate_sample_mean(NYC_payroll, 'Base Salary', 100, 12000) 
plots.xlim(50000, 100000);
```

Notice that our distribution of *sample means* looks approximately normal! Did it matter if our original distribution was normally distributed?

**Question 2.3.** Simulate two sample means, one for a sample of 300 salaries and one for a sample of 550 salaries. In each case, perform 12,000 repetitions. Don't worry about the `plots.xlim` or `plots.ylim` lines – they just make sure that all of the plots have the same x-axis and y-axis, respectively. 


```python
simulate_sample_mean(..., ..., ..., ...)
plots.xlim(50000, 100000)
plots.ylim(0, 0.0002)
plots.show();
print('\n')
simulate_sample_mean(..., ..., ..., ...)
plots.xlim(50000, 100000)
plots.ylim(0, 0.0002)
plots.show();
```

**Question 2.4.** Assign `q2_4` to an array of numbers corresponding to true statement(s) about the plots from 2.3. 

*Hint*: If a distribution is bell-shaped, what type of curve does it follow? 

1. The Central Limit Theorem doesn't apply because the population is not bell-shaped.
2. The Central Limit Theorem applies because even though the population distribution is skewed, the distribution of sample means appears  bell-shaped.
3. The sample mean of 550 salaries is likely to be closer to the population mean compared to the sample mean of 300 salaries according to the CLT. 
4. According to the CLT the mean of all the sample means usually equals the value of a particular sample mean.
5. The CLT requires we draw a large random sample with replacement.
6. The SD of the mean increases in value with larger sample sizes.


```python
q2_4 = ...
```

```python
grader.check("q2_4")
```

**Question 2.5.** Assign `q2_5` to an array of numbers corresponding to true statement(s) about the plots from 2.3.

1. The distribution of sample means for sample size 550 has **more** variability than the distribution of sample means for size 300.
2. The distribution of sample means for sample size 550 has **less** variability than the distribution of sample means for size 300.
3. The mean of the plots in 2.3 is approximately equal to the mean of one of the samples.
4. The mean of the plots in 2.3 is approximately equal to the population mean.


```python
q2_5 = ...
```

```python
grader.check("q2_5")
```

Below, we'll look at what happens when we take an **increasing number of resamples of a fixed sample size.** Notice what number in the code changes, and what stays the same. How does the distribution of the resampled means change?

```python
simulate_sample_mean(NYC_payroll, 'Base Salary', 100, 500)
plots.xlim(50000, 100000)
plots.ylim(0, 0.00016);
```

```python
simulate_sample_mean(NYC_payroll, 'Base Salary', 100, 1000)
plots.xlim(50000, 100000);
plots.ylim(0, 0.00015);
```

```python
simulate_sample_mean(NYC_payroll, 'Base Salary', 100, 5000)
plots.xlim(50000, 100000);
plots.ylim(0, 0.00015);
```

```python
simulate_sample_mean(NYC_payroll, 'Base Salary', 100, 10000)
plots.xlim(50000, 100000);
plots.ylim(0, 0.00015);
```

What did you notice about the distributions of sample means in the four histograms above? Discuss with your peers or ask a staff member.

**Question 2.6.** Assign the variable `SD_of_sample_means` to the integer corresponding to your answer to the following question:

When I increase the number of *resamples* that I take, for a *fixed* sample size, the SD of my sample means will...

1. Stay about the same
2. Decrease
3. Increase
4. Vary widely


```python
SD_of_sample_means = ...
```

**Question 2.7.** Let's think about how the relationships between population SD, sample SD, and SD of sample means change with varying sample size. Which of the following is true? Assign the variable `pop_vs_sample` to an array of integer(s) that correspond to true statement(s).

*Hint 1:* The sample SD is different from the SD of sample means. 

1. As the sample size increases, the SD of sample means increases.
2. As the sample size increases, the SD of sample means decreases.
3. As the sample size increases, the SD of sample means stays the same.
4. As the sample size increases, the sample SD increases.
5. As the sample size increases, the sample SD decreases.
6. As the sample size increases, the sample SD becomes more consistent with population SD.


```python
pop_vs_sample = ...
```

```python
grader.check("q2_7")
```

**Question 2.8.** Is there a relationship between the sample size and the standard deviation of the sample means? Assign `q2_8` to the number corresponding to the statement that answers this question.

*Hint:* [Chapter 14.5](https://inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html) of the textbook may be helpful.

1. There is no relationship between the sample size and the the SD of the sample means.
2. The SD of the sample means is equal to the population SD divided by the square of the sample size.
3. The SD of the sample means is equal to the population SD.
4. The SD of the sample means is equal to the population SD divided by the square root of sample size.


```python
q2_8 = ...
```

Throughout this lab, we have been taking many random samples from a population. However, all of these principles hold for bootstrapped resamples from a single sample. If your original sample is relatively large, all of your re-samples will also be relatively large, and so the SD of resampled means will be relatively small. 

In order to change the variability of your sample mean, you’d have to change the size of the original sample from which you are taking bootstrapped resamples.

## 3. Submission

<img src="jim.jpg" alt="drawing" width="300"/>

Jim really wanted to congratulate you on completing Lab 8!

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_3 (points: None)
```
>>> round(mean_age, 3) == 23.128
True
```
```
>>> round(sd_age, 3) == 4.640
True
```


Tests for q1_4 (points: None)
```
>>> round(overall_mean, 3) == 63.370
True
```
```
>>> round(overall_sd, 3) == 8.036
True
```


Tests for q2_1 (points: None)
```
>>> np.random.seed(8)
>>> np.round(get_sample_mean(NYC_payroll, 'Base Salary', 100), 3) == 60441.993
True
```


Tests for q2_4 (points: None)
```
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> hash_user = get_hash(set(np.array(q2_4).astype(int)))
>>> hash_expected = get_hash(set(np.array(make_array(2, 3, 5)).astype(int)))
>>> hash_user == hash_expected
True
```


Tests for q2_5 (points: None)
```
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> hash_user = get_hash(set(np.array(q2_5).astype(int)))
>>> hash_expected = get_hash(set(np.array(make_array(2, 4)).astype(int)))
>>> hash_user == hash_expected
True
```


Tests for q2_6 (points: None)


Tests for q2_7 (points: None)
```
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> hash_user = get_hash(set(np.array(pop_vs_sample).astype(int)))
>>> hash_expected = get_hash(set(np.array(make_array(2, 6)).astype(int)))
>>> hash_user == hash_expected
True
```


Tests for q2_8 (points: None)
