```python
# Initialize Otter
import otter
grader = otter.Notebook("lab06.ipynb")
```

## Lab 6: Take the Challenge! The Pepsi Challenge

Welcome to Lab 6!

After such an extensive introduction to programming for data science, we are finally moving into the section of the course where we can apply our new skils to answer real questions.  

In this lab, we'll use testing techniques that were introduced in lecture to test the idea of blind taste testing, and wheather we like products because of taste or marketing. 


**Submission**: Once you’re finished, run all cells besides the last one, select File > Save Notebook, and then execute the final cell. Then submit the downloaded zip file, that includes your notebook,  according to your instructor's directions.

```python
# Run this cell, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
from matplotlib import patches
from ipywidgets import interact, interactive, fixed
import ipywidgets as widgets
from IPython.display import YouTubeVideo 
import d8error
```

## 1. What is the Pepsi Challenge
The Pepsi Challenge was a direct marketing campaign by PepsiCo to compare the taste of Pepsi to that of its main competitor, Coca-Cola. In a blind taste test, participants were asked to sample two different cola drinks without knowing the brands and then select which one they preferred. The challenge's premise was simple yet effective: if a majority of participants preferred Pepsi over Coca-Cola in a blind taste test, it would suggest that Pepsi's taste was superior, challenging public perceptions and brand loyalty based on taste alone.


### The Origins of the Pepsi Challenge
[The Pepsi Challenge](https://en.wikipedia.org/wiki/Pepsi_Challenge)  first began in the 1970s as part of PepsiCo's aggressive marketing strategy to compete against Coca-Cola, the market leader. By conducting public blind taste tests, Pepsi aimed to demonstrate that consumers preferred the taste of Pepsi to Coca-Cola, thus encouraging people to rethink their cola preferences based on the taste rather than brand reputation or loyalty.


```python
YouTubeVideo("Qgr5kc5QdBU", height = 600, width = 550)
```

### The Experiment
A student, Ramona, wanted to replicate Pepsi's experiment, to find out if Coca-Cola fans can actually tell the difference between Pepsi and Coca-Cola by taste. Ramona was skeptical. She found 200 sworn Coca-Cola fans for our experiment. Participants were given two samples of cola, one of Pepsi and one of Coca-Cola, without any branding or identification. They then were asked to identify which one was Pepsi and which one was Coca-Cola. If a participant was truly a Coca-Cola fan, it would be expected that they could correctly identify the Coca-Cola sample.

In total, 200 taste tests were conducted, participants picked the correct cola 55% of the time.

Ramona's main goal here was to test whether or not the Coca-Cola fans' guesses were random, like the flip of a coin.

We will now begin to formulate this experiment in terms of the terminology we learned in this course.

**Question 1.1**:  Describe Ramona's model for how likely the Coca-Cola fans are to identify the correct cola. What alternative model is her model meant to discredit? Discuss with students around you to come to a conclusion. Check in with a TA you are stuck.


_Type your answer here, replacing this text._

**Question 1.2:** In Ramona's version of the Pepsi Challenge, participants identified the right cola 55% (.55) of the time.
According to her model, on average, what proportion of times does she expect participants to identify the correct cola? Provide your answer as a decimal between 0 and 1, reflecting the observed proportion in our experiment.


```python
expected_correct_prop = ...
expected_correct_prop
```

```python
grader.check("q1_2")
```

The goal now is to see if our deviation from this expected proportion of correct answers is due to something other than chance. 

**Question 1.3:** We usually use a statistic to help determine which model the evidence points towards. What is a statistic that we can use to compare outcomes under Ramona's model to what was observed? Assign `viable_stat` to an array of integer(s) representing test statistics that we can use: 

1. The absolute difference between the expected percent correct and the actual percent correct
2. The sum of the expected percent correct and the actual percent correct
3. The difference between the expected percent correct and the actual percent correct

**NOTE:** Make sure to use `make_array` to create your array of integer(s)!


```python
viable_stat = ...
viable_stat
```

```python
grader.check("q1_3")
```

**Question 1.4:** Why is the statistic selected in Question 1.3 the appropriate choice for comparing outcomes in the Pepsi Challenge experiment? Discuss how it relates to the models you defined in Question 1.1.


_Type your answer here, replacing this text._

**Question 1.5:** Define the function `stat` which takes in an expected proportion and an actual proportion, and returns the value of the statistic chosen in Question 1.3. Assume that the argument takes in proportions, but  return your answer as a percentage. 

*Hint:* Remember we are asking for a **percentage**, not a proportion. 


```python
def stat(expected_prop, actual_prop):
    ...
```

```python
grader.check("q1_5")
```

**Question 1.6:** Use your newly defined function to calculate the observed statistic from Pepsi's experiment. 


```python
observed_stat = ...
observed_stat
```

```python
grader.check("q1_6")
```

**Is this observed statistic consistent with what we expect to see under Ramona’s model?**

In order to answer this question, we must simulate the experiment as though Ramona's model was correct, and calculate our statistic for every simulation.

### `sample_proportions`

`sample_proportions` can be used to randomly sample from multiple categories when you know the proportion of data points that are expected to fall in each category. `sample_proportions` takes two arguments: the sample size and an array that contains the distribution of categories in the population (should sum to 1).

Consider flipping a fair coin, where the two outcomes (coin lands heads and coin lands tails) occur with an equal chance. We expect that half of all coin flips will land heads, and half of all coin flips will land tails.

Run the following cell to see the simulation of 20 flips of a fair coin. Let the first item of `coin_proportions` be the proportion of heads and the second item of `coin_proportions` be the proportion of tails.

*Observe what happens when you run this cell multiple times—the proportion of coin flips that land heads and tails appears to change, as you are simulating flipping 20 coins each time!*

```python
coin_proportions = make_array(0.5, 0.5) 
twenty_flips = sample_proportions(20, coin_proportions)
twenty_flips
```

`sample_proportions` returns an array that is the same length as the proportion array that is passed through. It contains the proportion of each category that appears in the sample. 

In our example, the first item of `twenty_flips` is the simulated proportion of heads and the second item of `twenty_flips` is the simulated proportion of tails.

```python
simulated_proportion_heads = twenty_flips.item(0)
simulated_proportion_tails = twenty_flips.item(1)

print("In our simulation, " + str(simulated_proportion_heads) + " of flips were heads and " \
      + str(simulated_proportion_tails) + " of flips were tails.")
```

**Question 1.7:** To begin simulating, we should start by creating a representation of Ramona's model to use for our simulation. This will be an array with two items in it. The first item should be the proportion of times a Coca-Cola fan identifies the correct cola, assuming that Ramona's model was correct. The second item should be the proportion of times, under the same assumption, that the Coca-Cola fan identifies the incorrect cola. Assign `model_proportions` to this array. 

After this, we can simulate 200 taste tests, and find a single statistic to summarize this instance of the simulation. Use the `sample_proportions` function and assign the **proportion of correct taste choices** (out of 200) to `simulation_correct_prop`. Lastly, use your statistic function to assign `one_stat`  to the value of the statistic for this one simulation.

*Hint:* `sample_proportions` usage can be found on the [Python Reference](http://data8.org/fa21/python-reference.html).


```python
# This saves the random state of our code so that we can 
# generate the same numbers each time we run the code.
# Please do not change this line. 
np.random.seed(16)

model_proportions = ...
simulation_correct_prop = ...
one_stat = ...
one_stat
```

```python
grader.check("q1_7")
```

**Question 1.8:** Let's now see what the distribution of statistics is actually like under Ramona's model. 

Define the function `simulation_and_stat` to take in the `model_proportions` array and the expected proportion of times a Coca-Cola fan would identify the correct cola under our model. The function should simulate our running through the experiment 200 times and return the statistic of this one simulation. 

*Hint:* This should follow the same pattern as the code you did in the previous problem.  

```python
def simulation_and_stat(model_proportions, expected_proportion_correct):
    '''Simulates 200 Pepsi and Coca-Cola fan taste test. 
    Returns one statistic from the simulation.'''
    # BEGIN SOLUTION
    simulated_proportion = sample_proportions(200, model_proportions).item(0)
    simulated_stat = stat(expected_proportion_correct, simulated_proportion)
    return simulated_stat
    # END SOLUTION
```

Using this function, assign `simulated_stats` to an array of 1500 statistics that you calculated under the assumption that Ramona's model was true.


```python
num_repetitions = 1500

simulated_stats = ...

for ... in ...:
    ...
```

```python
grader.check("q1_8")
```

Let's view the distribution of the simulated statistics under Ramona's model, and visually compare where the observed statistic lies relative to the simulated statistics.

```python
t = Table().with_column('Simulated Statistics', simulated_stats)
t.hist()
plt.scatter(observed_stat, 0, color='red', s=100, zorder=2);
```

We can make a visual argument as to whether we believe the observed statistic is consistent with Ramona's model. Here, since larger values of the test statistic suggest the alternative model (where the chance of guessing the correct cola is something other than 50%), we can formalize our analysis by finding what proportion of simulated statistics were as large or larger than our observed test statistic (the area at or to the right of the observed test statistic). If this area is small enough, we’ll declare that the observed data are inconsistent with our simulated model. Here is the [link](https://inferentialthinking.com/chapters/11/1/Assessing_a_Model.html) to the section in the textbook.

**Question 1.9:** Calculate the proportion of simulated statistics greater than or equal to the observed statistic. 

*Hint:* `np.count_nonzero` usage can be found [here](http://data8.org/fa21/python-reference.html).


```python
proportion_greater_or_equal = ...
proportion_greater_or_equal
```

```python
grader.check("q1_9")
```

By convention, we often compare the proportion we just calculated to 0.05. If the proportion of simulated stats greater than or equal to the observed stat is sufficiently small (less than or equal to 0.05), then this is evidence against Ramona's model. Conceptually, you may think of this as the case where less than 5% of simulated values are as far or farther away from what we had expected. If this is not the case, we don’t have any reason to doubt our model. 

This should help you make your own conclusions about our experiment. 

**Question 1.10:** Now, take some time to think to yourself and discuss with your peers:

1. Is the data more consistent with Ramona's model (Coca-Cola fans were randomly guessing)?
2. What does this mean in terms of Ramona's experiment? Do the Coca-Cola fans' answers follow an even chance model or is there something else at play? 

Did you talk to your peers? (T/F)


```python
peer_talk = ...
peer_talk
```

```python
peer_talk == True
```

## 2. Submission

<img src="appa_momo.jpg" alt="drawing" width="300"/>

Appa and Momo want to congratulate you on completing the lab!! Now you can relax like them! 

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_2 (points: None)
```
>>> 0 <= expected_correct_prop <= 1
True
```
```
>>> expected_correct_prop == 0.5
True
```


Tests for q1_3 (points: None)
```
>>> type(viable_stat) == np.ndarray
True
```
```
>>> int(sum(viable_stat))
1
```


Tests for q1_5 (points: None)
```
>>> int(round(stat(.5,.5) + stat(.4,.1),1))
30
```
```
>>> int(stat(.4,.1) - stat(.1,.4))
0
```


Tests for q1_6 (points: None)
```
>>> int(round(observed_stat,2))
5
```


Tests for q1_7 (points: None)
```
>>> len(model_proportions) % 2 == 0
True
```
```
>>> len(np.unique(model_proportions))
1
```
```
>>> sum(model_proportions) == 1
True
```
```
>>> type(simulation_correct_prop) == float
True
```
```
>>> round(simulation_correct_prop, 2)
0.49
```
```
>>> round(one_stat, 2)
1.0
```


Tests for q1_8 (points: None)
```
>>> len(simulated_stats) == 1500
True
```
```
>>> np.all(simulated_stats <= 30)
True
```
```
>>> np.all(simulated_stats >= 0)
True
```
```
>>> 0 <= simulation_and_stat(model_proportions, expected_correct_prop) <= 25
True
```


Tests for q1_9 (points: None)
```
>>> 0 <= proportion_greater_or_equal <= 1
True
```
```
>>> proportion_greater_or_equal*1500 == np.count_nonzero(simulated_stats >= observed_stat)
True
```
