```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

# Homework 5: Applying Functions and Iteration

Please complete this notebook by filling in the cells provided. Before you begin, execute the previous cell to load the provided tests.

**Helpful Resource:**
- [Python Reference](http://data8.org/sp22/python-reference.html): Cheat sheet of helpful array & table methods used in Data 8!

**Recommended Readings**: 

* [Tabular Thinking Guide](https://drive.google.com/file/d/13HaWu4rrjeowEOs6S_vaQGM3GremzXQ3/view?usp=sharing)
* [Applying Functions](https://www.inferentialthinking.com/chapters/08/1/Applying_a_Function_to_a_Column.html)
* [Conditionals](https://www.inferentialthinking.com/chapters/09/1/Conditional_Statements.html)
* [Iteration](https://www.inferentialthinking.com/chapters/09/2/Iteration.html)

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute this cell again.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. **Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Note: This homework has hidden tests on it. That means even though the tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**


Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

You should start early so that you have time to get help if you're stuck.

## 1. NBA Los Angeles Lakers Season

```python
# Run this cell to set up the notebook, but please don't change it.

# These lines import the numpy and datascience modules.
import numpy as np
from datascience import *
import d8error
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

Magic Johnson, former NBA player, and Laker legend, is trying to analyze how well the Los Angeles Lakers (LAL) performed from the 80s when he played to the 2010s when Kobe Bryant played. The table below has several columns detailing "Home Wins", "Home Loss", "Away Wins", "Away loss" of the history of the LAL from 1980 - 2010. The number of wins and losses that the Lakers have historically accumulated is stored in a table called `NBA Updated.csv`.

```python
#DO NOT EDIT
# Just run this cell
# Read in the "NBA Updated.csv" csv file
LAL_record = Table().read_table("NBA Updated.csv").where("Year", are.above(1979)).select("Year", "Team",\
"Home Wins", "Home Loss", "Away Wins", "Away Loss").where("Team", "Los Angeles Lakers").\
relabeled("Home Loss", "Home Losses").relabeled("Away Loss", "Away Losses")
LAL_record.show()
```

Each season, the LA Lakers accumulate wins from both home and away games, as well as losses from home and away games. Let's start by finding the total wins and losses for each season. 

**Question 1.** Write a function called `partial_record`.  It should take two arguments, each argument representing the number of home wins and away wins (or home losses and away losses) for the LA Lakers in a particular season. The function should return a partial record of wins (or losses) for that season. **(20 Points)**

*Note:* Don't overthink this question!

*Note:* Each row represents a season!

*Note:* It is partial record because you are only looking for total wins **OR** total losses.

```python
def partial_record(...):
    '''Returns the partial record calculated by adding up the home and away wins (or home and away losses) each year'''
    ...

partial_record(30, 27) #DO NOT CHANGE THIS LINE
```

```python
grader.check("q1_1")
```

**Question 2.** Create a new table `final_record` with three columns in this *specific* order: `Year`, `Wins`, and `Losses`. You will have to create the `Wins` and `Losses` columns. Use the function `partial_record` you just defined in the previous question for this problem. **(20 Points)**

*Hint:* If you want to apply a function that takes in multiple arguments, you can pass multiple column names as arguments in `tbl.apply()`. The column values will be passed into the corresponding arguments of the function. Take a look at the Python Reference Sheet and Lecture 13's demo for syntax.

*Note:* If you’re running into issues creating `final_record`, check that `Wins` and `Losses` output what you want. 


```python
lal_wins = LAL_record.apply(partial_record, "Home Wins", "Away Wins", )
lal_losses = LAL_record.apply(partial_record, "Home Losses", "Away Losses")
final_record = ...
final_record.show()
```

```python
grader.check("q1_2")
```

We can get specific row objects from a table. You can use `tbl.row(n)` to get the `n`th row of a table. `row.item("column_name")` will allow you to select the element that corresponds to `column_name` in a particular row. Here's an example:

```python
# Just run this cell
LAL_record.row(23).item("Away Wins")
```

```python

```

**Question 3.** We want to see for a particular season whether the LA Lakers had a losing season (more losses than wins). Write a function called `did_lal_lose`.  It should take one argument: a **row object** from the `final_record` table. It should return either `True` if LAL's losses were more than their wins, and `False` otherwise. **(20 Points)**

*Note 1*: "Row object" means a row from the table that contains all the data for that specific row. It is **not** the index of a row. Do not try and call `final_record.row(row)` inside of the function.


*Note 2*: If you still need clarification on row objects, try printing out final_record.row(1) in a new cell to see what it looks like visually! This piece of code pulls out the row object located at index 1 of the final_record table and returns it. When you display it in a cell, you'll see that it is not located within a table, but is instead a standalone row object!


```python
def did_lal_lose(row):
    ...

did_lal_lose(final_record.row(14)) #DO NOT CHANGE THIS LINE
```

```python
grader.check("q1_3")
```

**Question 4.** Magic Johnson wants to know the number of winning and losing seasons the Los Angeles Lakers had from 1980 to 2010. **Note: each year is a season.** Using the `final_record` table:

1. Assign `season_record` to an array of `True` and `False` values that correspond to whether or not LAL had a losing season.
2. Add the `season_record` array to the `final_record` table in a column named `Record`, and assign this to `final_record_with_results`.
3. Then, respectively assign the number of wins and losses LAL had to `lal_defeat` and `lal_victory`.

**(20 Points)**

*Hint 1*: `True` and `False` are **not** strings. What data type are they?

*Hint 2*: `tbl.num_rows` might be helpful too.

*Hint 3*: When you only pass a function name and no column labels through `tbl.apply()`, the function gets applied to every row in `tbl`.



```python
season_record = ...
final_record_with_results = ...
lal_defeat = ...
lal_victory = ...


#Don't delete or edit the following line:
print(f"Fom 1980 to 2010, the LA Lakers had {lal_victory} winning seasons and {lal_defeat} losing seasons. Go Laker Nation!")
```

```python
grader.check("q1_4")
```

**Question 5:** Sometimes in the NBA, teams have great seasons where they win most of their games. Other times, they have seasons where they have lost a majority of their games or barely won more games than lost. Let's define a **great season** to be a season in which the LA Lakers won more than 59 games.

Use your `final_record` table to assign `great_win_season` to an array to the years that LAL had more that 59 games won. You may find the `is_great_win_season` function defined below helpful to you! **(20 Points)**


```python
def is_great_win_season(row):
    '''Return a boolean to describe whether or not a season (row) is a great season'''
    games_won = row.item("Wins")
    
    if  games_won > 59:
        return True
    else:
        return False

great_win_season = make_array()

for row in final_record.rows: # This will help us iterate through rows of final_scores table
    great_season = ...
   
    if is_great_win_season(row):
        great_win_season = np.append(great_win_season, great_season)# Do not change this line - Adds character to the end of some_string
great_win_season
```

```python

```

```python
grader.check("q1_5")
```

You're done with the required section of Homework 5! Continue on to the optional section for some more practice with iterations and for loops.

Make sure to **run the submit cell located at the bottom of this notebook.**

## 2. Unrolling Loops (Optional)

**This section of HW5 is optional. Do it for your own practice, but it will not be incorporated into the final grading!**

"Unrolling" a `for` loop means to manually write out all the code that it executes.  The result is code that does the same thing as the loop, but without the structure of the loop.  For example, for the following loop:

    for num in np.arange(3):
        print("The number is", num)

The unrolled version would look like this:

    print("The number is", 0)
    print("The number is", 1)
    print("The number is", 2)


Unrolling a `for` loop is a great way to understand what the loop is doing during each step. In this exercise, you'll practice unrolling a `for` loop.


In the question below, write code that does the same thing as the given code, but with any `for` loops unrolled.  It's a good idea to run both your answer and the original code to verify that they do the same thing.  (Of course, if the code does something random, you'll get a different random outcome than the original code!)

**Optional Question 1.** Unroll the code below.

```python
for joke_iteration in np.arange(3):
    print("Knock, knock.")
    print("Who's there?")
    print("Banana.")
    print("Banana who?")
print("Knock, knock.")
print("Who's there?")
print("Orange.")
print("Orange who?")
print("Orange you glad I didn't say banana?")
```

```python
""" # BEGIN PROMPT
...
"""; # END PROMPT
# BEGIN SOLUTION NO PROMPT
print("Knock, knock.")
print("Who's there?")
print("Banana.")
print("Banana who?")
print("Knock, knock.")
print("Who's there?")
print("Banana.")
print("Banana who?")
print("Knock, knock.")
print("Who's there?")
print("Banana.")
print("Banana who?")
print("Knock, knock.")
print("Who's there?")
print("Orange.")
print("Orange who?")
print("Orange you glad I didn't say banana?")
# END SOLUTION
```

Congratulations, you are done with Homework 5!

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.** 

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [10, 10])
```
>>> partial_record(13,13)
26
```
```
>>> partial_record(22,12,)
34
```


Tests for q1_2 (points: [0, 0])
```
>>> final_record.num_columns
3
```
```
>>> set(['Year', 'Wins', 'Losses']) == set(final_record.labels)
True
```


Tests for q1_3 (points: [0, 10])
```
>>> 'did_lal_lose' in globals()
True
```
```
>>> did_lal_lose(final_record.row(1))
False
```


Tests for q1_4 (points: [0, 0])
```
>>> 0 <= lal_victory <= 30
True
```
```
>>> 0 <= lal_defeat <= 10
True
```


Tests for q1_5 (points: [5, 5])
```
>>> is_great_win_season(final_record.row(4)) is True
True
```
```
>>> # Make sure big_wins is an array.
>>> import numpy as np
>>> type(great_win_season) == np.ndarray
True
```
